```python
# Initialize Otter
import otter
grader = otter.Notebook("HW_06.ipynb")
```

```python
rng_seed = 42
```

# Homework 6

## K-Means Image Analysis, Neural Networks

## Problem 1: Flower Image Analysis
### 1) Objective

Before Convolutional Neural Networks (CNN) became available, K-means was an efficient method to segment images. Each pixel is represented by its RGB color vector. If we assume that pixel which belong to the same object should have the same color, then K-means should be able to cluster the pixel accordingly and therefore segment the image.<br> 
We want to analyze the Oxford 102 Flower Dataset using K-means clustering, hierarchical clustering, and decision tree classification and later in the course compare the results we obtain here to what we obtain using CNNs.<br>
<br>

### 2) Preparation

Before starting, we need to import the necessary libraries for data analysis and visualization.<br>
First, we import the standard libraries: 

```python
import numpy as np
import seaborn as sns
import matplotlib.pyplot as plt
from matplotlib.colors import rgb_to_hsv, hsv_to_rgb
from scipy.spatial.distance import pdist, squareform
```

Next, we need to load the dataset from `torchvision`

```python
from torchvision.datasets import Flowers102 #pip install torchvision if necessary
```

Finally, we import our clustering methods:

```python
from sklearn.cluster import KMeans
from sklearn.tree import DecisionTreeClassifier, plot_tree
from sklearn.metrics import confusion_matrix, ConfusionMatrixDisplay
```

<br>

The dataset from `torchvision` contains various flower images and their corresponding labels.<br>
Run the following lines (it might take a few minutes!).

```python
transforms = lambda x: np.array(x) / 255
dataset_train = Flowers102(root='Data', split='train', download=True, transform=transforms)
dataset_val = Flowers102(root='Data', split='val', download=True, transform=transforms)
dataset_test = Flowers102(root='Data', split='test', download=True, transform=transforms)
labels_train = np.array(dataset_train._labels)
labels_val = np.array(dataset_val._labels)
labels_test = np.array(dataset_test._labels)
```

```python
classes = np.array(["pink primrose", "hard-leaved pocket orchid", "canterbury bells", "sweet pea","english marigold","tiger lily", "moon orchid",
    "bird of paradise", "monkshood","globe thistle", "snapdragon","colt's foot","king protea","spear thistle","yellow iris","globe-flower",
    "purple coneflower", "peruvian lily", "balloon flower", "giant white arum lily","fire lily","pincushion flower", "fritillary", "red ginger",
    "grape hyacinth","corn poppy", "prince of wales feathers","stemless gentian","artichoke","sweet william","carnation", "garden phlox",
    "love in the mist","mexican aster", "alpine sea holly", "ruby-lipped cattleya", "cape flower", "great masterwort", "siam tulip",  "lenten rose",
    "barbeton daisy", "daffodil", "sword lily", "poinsettia", "bolero deep blue", "wallflower","marigold", "buttercup", "oxeye daisy",
    "common dandelion", "petunia", "wild pansy","primula", "sunflower","pelargonium", "bishop of llandaff", "gaura", "geranium", "orange dahlia",
    "pink-yellow dahlia?","cautleya spicata", "japanese anemone", "black-eyed susan", "silverbush", "californian poppy", "osteospermum",
    "spring crocus", "bearded iris", "windflower", "tree poppy", "gazania", "azalea", "water lily","rose", "thorn apple", "morning glory",
    "passion flower", "lotus", "toad lily", "anthurium", "frangipani", "clematis", "hibiscus", "columbine", "desert-rose", "tree mallow", "magnolia",
    "cyclamen",  "watercress", "canna lily", "hippeastrum", "bee balm", "ball moss", "foxglove", "bougainvillea","camellia","mallow","mexican petunia",
    "bromelia", "blanket flower", "trumpet creeper", "blackberry lily"])
```

<!-- BEGIN QUESTION -->

<br>

2.1) We can display images in python using `plt.imshow(x)`, where *x* is the image. Show the first image from the dataset to understand what the data looks like.

```python
index = 0 #first image
x, y = dataset_train[index] #x: image, y: the label, classes[y]: the actual name of the flower (type str) 

######## Fill in the code below ########
...
########################################
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<br>

2.2) Print the shape and type of the image. Generate three plots, each showing the image only in one color channel in gray scale (`cmap='gray'`)! 

```python
######## Fill in the code below ########
...
########################################
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<br>

2.3) Show an image in the first 8 classes from the dataset to understand what the data looks like.

```python
indices = [0, 10, 20, 30, 40, 50, 60, 70]

######## Fill in the code below ########
...
########################################
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<br>

### 3) Image Segmentation

3.1) Segment the first image in the training set using K-means clustering. Reshape the image into a 2D array where each row represents a pixel and each column represents a color channel (RGB), and then apply K-means clustering to cluster the pixels. Visualize the segmented image and discuss what you observe.

```python
index = 0

######## Fill in the code below ########
...
########################################
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

> The segmented image shows distinct color regions corresponding to different parts of the flower and background. Visually, the first cluster represents the background, the second cluster represents the flower, and the third cluster represents the stems. This segmentation helps in identifying different components of the image based on color similarity.

3.2) Visualize the center of each cluster as a color using `plt.imshow(colors.reshape(.., .., ..))`. What does this tell you about the image?

```python
######## Fill in the code below ########
...
########################################
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

> The cluster centers represent the dominant colors in the image. The first cluster center is a shade of dark green corresponding to the ground, the second cluster center is a shade of pinkish purple corresponding to the flower, and the third cluster center is a shade of brown corresponding to the stems. This indicates that the image has distinct color regions that can be effectively captured using K-means clustering.

3.3) Reconstruct the image using the cluster centers and the labels assigned to each pixel. Does the reconstructed image capture the main features of the original image? Discuss your observations.

_Type your answer here, replacing this text._

```python
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### 4) Image Clustering

4.1) Use hierarchical clustering (`sns.clustermap`) to cluster the first 8 classes of images in the training set. Compute the mean color of the first image in the class and use the mean colors as features. Next, based on the mean colors, generate a distance square matrix (you can use `pdist(your_means, metric='euclidean')`) for clustering. Visualize the resulting dendrogram and discuss what you observe.

_Type your answer here, replacing this text._

```python
indices = [0, 10, 20, 30, 40, 50, 60, 70]
ticks = classes[labels_train[indices]]

######## Fill in the code below ########
...
########################################
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<br>

### 5) Classification

5.1) Perform decision tree classification using `DecisionTreeClassifier` to predict the flower species based on the features from the K-means clustering. Train the decision tree classifier for class 0 (pink primrose) and class 4 (english marigold) on the training set and evaluate its performance on the validation set by creating a **confusion matrix plot** for both, training and evaluation set. Discuss why or why not the features from K-means clustering might be effective for this classification task.

_Type your answer here, replacing this text._

```python
indices_train = np.array([0, 1, 2, 3, 4, 5, 6, 7, 8, 9, 40, 41, 42, 43, 44, 45, 46, 47, 48, 49])
indices_val = np.array([0, 1, 2, 3, 4, 5, 6, 7, 8, 9, 40, 41, 42, 43, 44, 45, 46, 47, 48, 49])

######## Fill in the code below ########
def extract_features(image, n_clusters=2):
    ...

...

tree = DecisionTreeClassifier(random_state=rng_seed) 
tree.fit(x_train, y_train)

...
########################################
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

5.2) Visualize the decision tree and interpret the results. Discuss which features are most important for classification and how well the model performs. You can visualize the features as colors to help interpret the decision tree.

_Type your answer here, replacing this text._

```python
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

5.3) Try a little more feature engineering by, for example, converting the RGB color space to HSV color space (`image = rgb_to_hsv(image)`) before applying K-means clustering, as well as sorting the cluster centers before flattening them into a feature vector. See if these changes improve the classification performance.

_Type your answer here, replacing this text._

```python
indices_train = np.array([0, 1, 2, 3, 4, 5, 6, 7, 8, 9, 40, 41, 42, 43, 44, 45, 46, 47, 48, 49])
indices_val = np.array([0, 1, 2, 3, 4, 5, 6, 7, 8, 9, 40, 41, 42, 43, 44, 45, 46, 47, 48, 49])

######## Fill in the code below ########
...
########################################
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

5.4) Try your final classification method on class 0 (pink primrose) and class 1 (hard-leaved pocket orchid). How well does it perform? Discuss what kind of features we have not considered that might help improve the classification.

_Type your answer here, replacing this text._

```python
indices_train = np.array([0, 1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19])
indices_val = np.array([0, 1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19])

######## Fill in the code below ########
...
########################################
```

<!-- END QUESTION -->


## Problem 2: Two Moons Classification with scikit-learn’s `MLPClassifier`

In this problem, you’ll use **scikit-learn’s built-in multilayer perceptron (MLP)** to classify the famous **two-moons** dataset.

You’ll explore how the architecture (number of hidden layers and neurons) and regularization affect decision boundaries and overfitting.

---

**What you’ll practice**
- Using `MLPClassifier` from `sklearn.neural_network`
- Controlling hyperparameters such as hidden-layer size, activation, solver, learning rate, and regularization
- Visualizing loss curves and decision boundaries
- Interpreting underfitting and overfitting behavior

---

**Rules**
- You may use any parts of scikit-learn and matplotlib.
- Do not use TensorFlow, PyTorch, or other deep learning frameworks.
- Keep everything in this notebook.


```python

import numpy as np
import matplotlib.pyplot as plt
from sklearn.datasets import make_moons
from sklearn.model_selection import train_test_split
from sklearn.neural_network import MLPClassifier
from sklearn.metrics import accuracy_score

np.random.seed(rng_seed)
```

<!-- BEGIN QUESTION -->

### 1) Generate and visualize the two-moons dataset

The dataset is generated through make_moons. You will need to split the dataset using train_test_split, with a 0.3 test size and setting the random_state to rng_seed. Plot a figure of your training data, with a different colour for each class.

```python

X, y = make_moons(n_samples=1000, noise=0.25, random_state=rng_seed)
X_train, X_val, y_train, y_val = ...

...
```

<!-- END QUESTION -->

### 2) Train an `MLPClassifier`

As you can see from above, the dataset generated is hard to classify with simpler methods such as k-means clustering or a linear regression. However, a simple Neural Network would be perfect for this task.

Make use of the MLPClassifier class in sklearn, start with a small network, train and classify on the training sample. Experiment with different parameters in the network such as hidden_layer_sizes, activation, solver, alpha (regularization strength) and learning rates. To pass the test, you need a certain validation accuracy that would match with what we would expect



```python
# Define and train MLP
mlp = ...
...

# Training results
print(f"Training accuracy: {mlp.score(X_train, y_train):.3f}")
print(f"Validation accuracy: {mlp.score(X_val, y_val):.3f}")
```

```python
grader.check("q2.2")
```

<!-- BEGIN QUESTION -->


### 3) Inspect the training loss curve
The attribute `mlp.loss_curve_` stores the loss per iteration. Plot the training loss curve


```python
...
```

<!-- END QUESTION -->


### 4) Cross-Validation of parameters

As in many other ML models, finding the best set of model hyper-parameters is an important challenge. This becomes unfeasible for extremely large networks due to the large hyper-parameter space, but in our case here, it still remains very possible. Let us try to search for the optimal parameters for our model:

Use `GridSearchCV` on a pipeline `Pipeline(StandardScaler, MLPClassifier(max_iter=700, random_state=rng_seed))` with:
- `clf__hidden_layer_sizes`: `(8,), (32,), (64,), (32,32)`
- `clf__alpha`: `1e-5, 1e-3, 1e-1`
- `clf__activation`: `'relu', 'tanh'`

Fit on **`X_train, y_train`** using `StratifiedKFold(n_splits=5, shuffle=True, random_state=rng_seed)`.  
Store the fitted grid search in **`gs`** and the best estimator in **`best`**.

```python
from sklearn.model_selection import GridSearchCV, StratifiedKFold
import warnings
from sklearn.exceptions import ConvergenceWarning

warnings.simplefilter("ignore", category=ConvergenceWarning)
```

```python
param_grid = {
    "hidden_layer_sizes": [(8,), (32,), (64,), (32,32)],
    "alpha": [1e-5, 1e-3, 1e-1],
    "activation": ["relu", "tanh"],
}
mlp_cv = ...

cv = ...
gs = ...
gs.fit(X_train, y_train)

best = gs.best_estimator_
print("Best params:", gs.best_params_)
print("CV best score:", gs.best_score_)
print("Best model val acc:", best.score(X_val, y_val))
```

```python
grader.check("q2.4")
```

<!-- BEGIN QUESTION -->


### 5) Visualize the decision boundary

Write some code to visiualize the decision boundary using your best MLPClassifier. Hint: you can do this by first making a np.meshgrid of values for x1, x2 and classifying every point on the grid with a probability. You can then use a contourf plot to plot the result.

```python
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### 6) Confusion Matrix

Plot the confusion matrix on the train and prediction data, using the best model from before.

```python
print('Training accuracy:', best.score(X_train, y_train))
cm = ...
...

print('Validation accuracy:', best.score(X_val, y_val))
cm = ...
...
```

<!-- END QUESTION -->



## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit the zip file to Gradescope!

```python
# Save your notebook first, then run this cell to export your submission.
grader.export()
```

Tests for q2.2 (points: 2)
```
>>> tr2, va2 = (mlp.score(X_train, y_train), mlp.score(X_val, y_val))
>>> assert va2 >= 0.92, f'Expected validation accuracy >= 0.92; got {va2:.3f}.'
>>> print('q2.2 public tests passed.')
q2.2 public tests passed.
```


Tests for q2.4 (points: 3)
```
>>> assert 'gs' in globals() and 'best' in globals(), 'Missing `gs` or `best`.'
>>> from sklearn.model_selection import GridSearchCV as _GSCV
>>> assert isinstance(gs, _GSCV), '`gs` should be a GridSearchCV object.'
>>> va_best = best.score(X_val, y_val)
>>> assert va_best >= 0.925, f'Best model validation accuracy should be at least 0.925; got {va_best:.3f}.'
>>> print('q2.4 public tests passed.')
q2.4 public tests passed.
```
